# BSTree

## Overview

In this notebook you will modify the binary search tree implementation completed in class so that it can be used as a mapping structure. The `Node` class will be updated so as to hold separate key and value attributes (instead of a single value, as it currently does), and instead of the `add` method, you should implement the [`__getitem__`](https://docs.python.org/3.5/reference/datamodel.html#object.__getitem__) and [`__setitem__`](https://docs.python.org/3.5/reference/datamodel.html#object.__setitem__) methods in order to associate keys and values. `__delitem__`, `__contains__`, and `__iter__` will also need to be updated, to perform key-based removal, search, and iteration. Finally, the `keys`, `values`, and `items` methods will return iterators that allow the keys, values, and key/value tuples of the tree (all sorted in order of their associated keys) to be traversed.

If `__setitem__` is called with an existing key, the method will simply locate the associated node and update its value with the newly provided value (as you would expect a mapping structure to do). If either `__getitem__` or `__delitem__` are called with a key that does not exist in the tree, a `KeyError` should be raised.

The API described above will allow the tree to be used as follows:

    t = BSTree()
    t[0] = 'zero'
    t[5] = 'five'
    t[2] = 'two'

    print(t[5])
    
    t[5] = 'FIVE!!!'

    for k,v in t.items():
        print(k, '=', v)

    del t[2]

    print('length =', len(t))
    
The expected output of the above follows:

    five
    0 = zero
    2 = two
    5 = FIVE!!!
    length = 2

The following `BSTree` class contains an updated `Node`, and stubs for the methods you are to implement. The first few simple test cases beneath the class definition should help clarify the required behavior.

<!-- grade_id: bstree, points: 0, grade: false, solution: true -->
```python
#<GRADED>
class BSTree:
    class Node:
        def __init__(self, key, val, left=None, right=None):
            self.key = key
            self.val = val
            self.left = left
            self.right = right
            
    def __init__(self):
        self.size = 0
        self.root = None
        
    def __getitem__(self, key):
        if not self.root: 
            return None
        
         # start at the root
        node = self.root
        while node:
            if node.key == key:
                return node.val
            elif node.left == None and node.right == None:
                raise KeyError
            elif key < node.key:
                # You go left if the given key is less-than the node's key.  
                node = node.left
            elif key > node.key:
                # You go right if the key is greater-than the node's key.
                node = node.right
            else:
                raise KeyError
        return None
    
    def __setitem__(self, key, val):
        if self.root == None:
            self.root = BSTree.Node(key, val)
            self.size +=1
        else:
            node = self.root

            while node:
                if node.key == key:
                    node.val = val
                    break
                elif node.right == None and node.left == None:
                    if key < node.key: # if there are no more children then insert new key and val as child
                        node.left = BSTree.Node(key, val)
                        self.size +=1
                    else:
                        node.right = BSTree.Node(key, val)
                        self.size +=1
                elif key < node.key:
                    # You go left if the given key is less-than the node's key.  
                    node = node.left
                elif key > node.key:
                    # You go right if the key is greater-than the node's key.
                    node = node.right
                
    
    def __delitem__(self, key):
        def del_rec(t):
            if not t: #empty spot in the tree, nothing to do
                return None
            elif t.key < key:
                t.right = del_rec(t.right)
                return t
            elif t.key > key:
                t.left = del_rec(t.left)
                return t
            else: #t.key == key
                if not t.left and not t.right:
                    return None
                elif not t.left and t.right:
                    return t.right
                elif not t.right and t.left:
                    return t.left

        assert(key in self)
        #deal with relatively simple case first!
        self.root = del_rec(self.root)
        self.size -=1
        
    def __contains__(self, key):
        node = self.root
        while node:
            if node.key == key:
                return True
            elif node.left == None and node.right == None:
                raise KeyError
            elif key < node.key:
                # You go left if the given key is less-than the node's key.  
                node = node.left
            elif key >= node.key:
                # You go right if the key is greater-than the node's key.
                node = node.right
            else:
                raise KeyError
        return None
    
    def __len__(self):
        return self.size
    
    def __iter__(self):
        def iter_rec(t):
            if t:
                yield from iter_rec(t.left)
                yield t.key
                yield from iter_rec(t.right) #does the same thing as the for loop

        yield from iter_rec(self.root)
    
    def keys(self):
        for x in self:
            yield x
            
    def values(self):
        for x in self:
            yield self.__getitem__(x)

    def items(self): 
        for x, y in zip(self.keys(), self.values()):
            yield (x, y)
        
    def pprint(self, width=64):
        """Attempts to pretty-print this tree's contents."""
        height = self.height()
        nodes  = [(self.root, 0)]
        prev_level = 0
        repr_str = ''
        while nodes:
            n,level = nodes.pop(0)
            if prev_level != level:
                prev_level = level
                repr_str += '\n'
            if not n:
                if level < height-1:
                    nodes.extend([(None, level+1), (None, level+1)])
                repr_str += '{val:^{width}}'.format(val='-', width=width//2**level)
            elif n:
                if n.left or level < height-1:
                    nodes.append((n.left, level+1))
                if n.right or level < height-1:
                    nodes.append((n.right, level+1))
                repr_str += '{val:^{width}}'.format(val=n.key, width=width//2**level)
        print(repr_str)
    
    def height(self):
        """Returns the height of the longest branch of the tree."""
        def height_rec(t):
            if not t:
                return 0
            else:
                return max(1+height_rec(t.left), 1+height_rec(t.right))
        return height_rec(self.root)
#</GRADED>
```

```python
t = BSTree()
0 in t
t[0] = 'zero'
print (t[0])
0 in t

t = BSTree()
key_vals = [(0, 'zero'), (2, 'two'), (1, 'one')]
sorted_key_vals = sorted(key_vals)

for k,v in key_vals:
    t[k] = v
    
for i,k in enumerate(t.keys()):
    k, sorted_key_vals[i][0]
    
    
```

Output:
```
zero
```

<!-- grade_id: simple_tests_1, points: 2, grade: true, solution: false -->
```python
# 2 points

from unittest import TestCase

tc = TestCase()
t = BSTree()
tc.assertEqual(len(t), 0)
tc.assertFalse(0 in t)
t[0] = 'zero'
tc.assertTrue(0 in t)
tc.assertEqual(len(t), 1)
```

<!-- grade_id: simple_tests_2, points: 2, grade: true, solution: false -->
```python
# 2 points

from unittest import TestCase

tc = TestCase()
t = BSTree()
tc.assertEqual(len(t), 0)
t[0] = 'zero'
tc.assertEqual(t[0], 'zero')
```

<!-- grade_id: simple_tests_3, points: 2, grade: true, solution: false -->
```python
# 2 points

from unittest import TestCase

tc = TestCase()
t = BSTree()
tc.assertEqual(len(t), 0)
t[0] = 'zero'
del t[0]
tc.assertFalse(0 in t)
tc.assertEqual(len(t), 0)
```

<!-- grade_id: simple_tests_4, points: 2, grade: true, solution: false -->
```python
# 2 points

from unittest import TestCase

tc = TestCase()
t = BSTree()
key_vals = [(0, 'zero'), (2, 'two'), (1, 'one')]
sorted_key_vals = sorted(key_vals)

for k,v in key_vals:
    t[k] = v

for i,k in enumerate(t.keys()):
    tc.assertEqual(k, sorted_key_vals[i][0])
```

<!-- grade_id: simple_tests_5, points: 1, grade: true, solution: false -->
```python
# 1 point

from unittest import TestCase

tc = TestCase()
t = BSTree()
key_vals = [(0, 'zero'), (2, 'two'), (1, 'one')]
sorted_key_vals = sorted(key_vals)

for k,v in key_vals:
    t[k] = v

for i,v in enumerate(t.values()):
    tc.assertEqual(v, sorted_key_vals[i][1])
```

<!-- grade_id: simple_tests_6, points: 1, grade: true, solution: false -->
```python
# 1 point

from unittest import TestCase

tc = TestCase()
t = BSTree()
key_vals = [(0, 'zero'), (2, 'two'), (1, 'one')]
sorted_key_vals = sorted(key_vals)

for k,v in key_vals:
    t[k] = v

for i,(k,v) in enumerate(t.items()):
    tc.assertEqual(k, sorted_key_vals[i][0])
    tc.assertEqual(v, sorted_key_vals[i][1])
```

<!-- grade_id: stress_testing, points: 5, grade: true, solution: false -->
```python
# 5 points

from unittest import TestCase
import random

tc = TestCase()
t = BSTree()
keys = list(range(100, 1000, 11))
random.shuffle(keys)
vals = [random.randrange(1000) for _ in range(100, 1000, 11)]

for i in range(len(keys)):
    t[keys[i]] = vals[i]

for i in range(len(keys)):
    tc.assertEqual(t[keys[i]], vals[i])
```

Output:
```
~\Anaconda3\lib\unittest\case.py in assertEqual(self, first, second, msg)
    827         """
    828         assertion_func = self._getAssertEqualityFunc(first, second)
--> 829         assertion_func(first, second, msg=msg)
    830 
    831     def assertNotEqual(self, first, second, msg=None):

~\Anaconda3\lib\unittest\case.py in _baseAssertEqual(self, first, second, msg)
    820             standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    821             msg = self._formatMessage(msg, standardMsg)
--> 822             raise self.failureException(msg)
    823 
    824     def assertEqual(self, first, second, msg=None):

AssertionError: None != 772
```

<!-- grade_id: iter_testing, points: 5, grade: true, solution: false -->
```python
# 5 points

from unittest import TestCase
import random

tc = TestCase()
t = BSTree()
keys = list(range(100, 1000, 11))
shuffled_keys = keys.copy()
random.shuffle(shuffled_keys)

for k in shuffled_keys:
    t[k] = str(k)

for i,k in enumerate(t.keys()):
    tc.assertEqual(k, keys[i])

for i,v in enumerate(t.values()):
    tc.assertEqual(v, str(keys[i]))

for i,(k,v) in enumerate(t.items()):
    tc.assertEqual(k, keys[i])
    tc.assertEqual(v, str(keys[i]))
```

Output:
```
~\Anaconda3\lib\unittest\case.py in assertEqual(self, first, second, msg)
    827         """
    828         assertion_func = self._getAssertEqualityFunc(first, second)
--> 829         assertion_func(first, second, msg=msg)
    830 
    831     def assertNotEqual(self, first, second, msg=None):

~\Anaconda3\lib\unittest\case.py in _baseAssertEqual(self, first, second, msg)
    820             standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    821             msg = self._formatMessage(msg, standardMsg)
--> 822             raise self.failureException(msg)
    823 
    824     def assertEqual(self, first, second, msg=None):

AssertionError: 188 != 100
```

<!-- grade_id: keyerror_tests, points: 5, grade: true, solution: false -->
```python
# 5 points

from unittest import TestCase
import random

tc = TestCase()
t = BSTree()
keys = list(range(0, 100, 2))
random.shuffle(keys)

for x in keys:
    t[x] = x*2

for k in range(1, 101, 2):
    with tc.assertRaises(KeyError):
        v = t[k]
```

Output:
```
~\Anaconda3\lib\unittest\case.py in __exit__(self, exc_type, exc_value, tb)
    201                                                                 self.obj_name))
    202             else:
--> 203                 self._raiseFailure("{} not raised".format(exc_name))
    204         else:
    205             traceback.clear_frames(tb)

~\Anaconda3\lib\unittest\case.py in _raiseFailure(self, standardMsg)
    133     def _raiseFailure(self, standardMsg):
    134         msg = self.test_case._formatMessage(self.msg, standardMsg)
--> 135         raise self.test_case.failureException(msg)
    136 
    137 class _AssertRaisesBaseContext(_BaseTestCaseContext):

AssertionError: KeyError not raised
```